A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Using Shell Commands from Jupyter Notebooks

In this problem, we will use IPython's ! function ([IPython as a system shell](http://ipython.readthedocs.io/en/stable/interactive/shell.html)) to perform some basic Unix shell commands.

```python
import os
from nose.tools import assert_equal, assert_true
```

You can use IPython/Jupyter notebook environment not only for Python itself, but for Unix shell commands and more. To run any shell commands from the notebook, simply prefix a shell command with !.

For example, the equivalent of `print()` for Unix is `echo`. So, to print out `Hello` in a terminal, we would first open a terminal by going to _New_ -> _Terminal_ on Jupyterhub. This gives us a command prompt:

```shell
data_scientist@jupyterhub:~$
```

I'll abbreviate the command prompt by simply writing `$`. At the command prompt, we would enter
``` shell
$ echo "Hello"
```
```
Hello
```

If we want to run the same shell command from the notebook, we prefix it with !.

```python
!echo "Hello"
```

You can also capture the results of a shell command and assign it to a Python variable with the following syntax:

```python
output = !echo "Hello"
```

Here, `output` is a Python variable (a list of strings for each line in the output). The standard output from `echo "Hello"` is captured and stored in the `output` variable. Let's print out the value of `output` using Python:

```python
print(output)
```

We see that the output of `echo` is stored as a list of strings.

In this notebook, we will use system shell commands to perform Unix data processing. I recommend that you solve the problem in the terminal first (because shell commands are meant to be executed in a terminal). Once you have an answer, repeat the same command in the notebook (don't forget to prefix the command with !). For example, if the question was

```
You are currently at /home/data_scientist. How do you change into /home/data_scientist/data?
```

the answer would be
```shell
$ cd /home/data_scientist/data
```

or

```shell
$ cd data
```

So, I would try it in the terminal and verify that the command works. After I'm convinced that I have the correct answer, I would return to the notebook, and start by writing ! in the code cell and type my answer.

```
!cd data
```

If you are comfortable with running shell commands directly from the notebook, you don't have to use the terminal, and you can directly enter your answers in the notebook. Your goal is to fill the blank code cells so that the tests (the assertion statements) don't produce any errors.

## Create a directory

Now that you understand how the shell commands work in IPython notebooks, let's practice what you have learned. In the following code cell,

- Use Unix shell commands to create a new directory at `/home/data_scientist/temp/data`.

<!-- grade_id: mkdir_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

The following test checks if you have successfully created the `data` directory.

<!-- grade_id: mkdir_test, points: 4, grade: true, solution: false -->
```python
assert_true(os.path.exists("/home/data_scientist/temp/data"))
```

Now, in `/home/data_scientist/data`, we have a file named `airports.csv`:

```shell
$ ls /home/data_scientist/data
```
```
2001	      airports.csv  email	     nltk_data	     weather
2001.csv      carriers.csv  misc	     plane-data.csv
2001.csv.bz2  delta.csv     ml-latest-small  textdm
```

```python
!ls /home/data_scientist/data
```

## Copy airports.csv

- Copy this file, `/home/data_scientist/data/airports.csv`, to the newly created directory, `/home/data_scientist/temp/data`.

<!-- grade_id: cp_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cp_test, points: 3, grade: true, solution: false -->
```python
assert_true(os.path.exists("/home/data_scientist/temp/data/airports.csv"))
```

## List the contents

- Verify that the file is now in the newly created directory by listing the contents of `/home/data_scientist/data/airports.csv`.
- Capture the results of your shell command, and assign it to a Python variable named `contents`. That is, you answer should begin with

```
contents = ! # YOUR CODE HERE
```

<!-- grade_id: contents_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

```python
print(contents)
```

<!-- grade_id: contents_test, points: 3, grade: true, solution: false -->
```python
assert_true(["airports.csv" in i for i in contents])
```

We don't need this temporary directory anymore, so we use `rm` to clean up.

```python
! rm -rf /home/data_scientist/temp
```