A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Plotting


In this problem, we use the matplotlib library to create a plot of the sine function.

Here's a sample plot.

![](sine.png)

You don't have to make your plot look exactly like this example. If your plot looks visually OK, and if the test code cell doesn't produce any errors, your solution is correct.

```python
%matplotlib inline

# plotting tools
import matplotlib as mpl
import matplotlib.pyplot as plt
import numpy as np
import seaborn as sns

# testing tools
from nose.tools import (
    assert_equal, assert_is_instance, assert_is_not,
    assert_greater_equal, assert_less_equal
)
from numpy.testing import assert_array_almost_equal
```

## Write a function named `plot_sine` that creates a plot of the sine function $\sin(x)$.

- Note that code for creating a `Figure` object and an `Axes` object is already provided:
```python
fig, ax = plt.subplots()
```
Also note that the `plot_sine` function returns an instance (named `ax`) of the `Axes` object. You should use `ax` to create your plot (check out the [lesson notebook on Python Plotting](https://datascience.business.illinois.edu/user/jkim575/notebooks/accy570-fa16_RO/Week3/notebooks/intro2plotting.ipynb) for examples on how to do this). The reason we write our function this way to return an `Axes` instance is because it can be used to extract data values for testing.

- Your plot should have a title and axis labels.

- The $x$-axis should be in radians (not degrees) and should at least go from $-\pi$ to $+\pi$. It is acceptable to show $\sin(x)$ for $x < -\pi$ or $x > \pi$.

<!-- grade_id: plot_sine_answer, points: 0, grade: false, solution: true -->
```python
def plot_sine():
    """
    Plots the sine function.
    
    Parameters
    ----------
    None
    
    Returns
    -------
    A matplotlib.Axes instance.
    """
    
    fig, ax = plt.subplots()

    # YOUR CODE HERE
    
    return ax
```

```python
ax = plot_sine()
```

<!-- grade_id: plot_sine_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(
    ax, mpl.axes.Axes,
    msg="Your function should return a matplotlib.axes.Axes object."
)

assert_equal(len(ax.lines), 1)

assert_is_not(
    len(ax.title.get_text()), 0,
    msg="Your plot doesn't have a title."
)
assert_is_not(
    ax.xaxis.get_label_text(), '',
    msg="Change the x-axis label to something more descriptive."
)
assert_is_not(
    ax.yaxis.get_label_text(), '',
    msg="Change the y-axis label to something more descriptive."
)
    
x, y = ax.lines[0].get_xydata().T
assert_less_equal(x.min(), -3.14)
assert_greater_equal(x.max(), 3.14)
assert_array_almost_equal(y, np.sin(x))
```