A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 2. Pandas_Database_Programming

In this problem, we execute SQL statements in Pandas to create and read a database.

```python
import os
import sqlite3 as sl
import pandas as pd

from nose.tools import assert_equal, assert_true
from pandas.util.testing import assert_frame_equal
```

Suppose we are given the following table.

<div align="center">
  <b>Midterm</b>
</div>
<table>
  <tr>
    <th>Id</th>
    <th>Name</th> 
    <th>Score</th>
  </tr>
  <tr>
    <td>1</td>
    <td>Alice</td> 
    <td>97.3</td>
  </tr>
  <tr>
    <td>2</td>
    <td>Bob</td> 
    <td>87.7</td>
  </tr>
  <tr>
    <td>3</td>
    <td>Chris</td> 
    <td>91.5</td>
  </tr>
</table>

In the following code cell, we represent this table as a Pandas DataFrame named `midterm`.

```python
>>> print(midterm)
```

```
     Name  Score
Id              
1   Alice   97.3
2     Bob   87.7
3   Chris   91.5
```

```python
midterm = pd.DataFrame(
    {
        "Id": [1, 2, 3],
        "Name": ["Alice", "Bob", "Chris"],
        "Score": [97.3, 87.7, 91.5]
    }
)

index = [1, 2, 3]
data = {
    "Name": ["Alice", "Bob", "Chris"],
    "Score": [97.3, 87.7, 91.5]
}
midterm = pd.DataFrame(data=data, index=index)
midterm.index.name = "Id"

print(midterm)
```

## Save DataFrame to SQL dabase

- Write a function that uses the `to_sql()` method in Pandas to write an existing Pandas DataFrame (`df`) to an SQL database file.
- `save_midterm` takes three arguments: `df`, `filename`, and `tablename`.
- The first paramter `df` is a Pandas DataFrame.
- The `filename` parameter is a string and specifies the name of the database file, e.g., `students.db`.
- The `tablename` parameter is a string and specifies the name of the SQL table, e.g. `Midterm`.
- Running
  ```python
  >>> save_midterm(midterm, "students.db", "Midterm")
  ```
  should take an existing Pandas DataFrame named `midterm`, and insert this data into an SQL table named `Midterm` in `students.db`. So, after running `save_midterm()` as shown above, we will have a file named `students.db` in the current working directory, and we can use what we learned in the previous week to query the database (see [ACCY 570 Week 12 Problem 1](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week12/assignments/Problem_1_SQL_Data_Manipulation.ipynb)). For example,
  ```python
  >>> !sqlite3 students.db ".schema"
  ```
  ```
  CREATE TABLE "Midterm" (
  "Id" INTEGER,
  "Name" TEXT,
  "Score" REAL
  );
  CREATE INDEX "ix_Midterm_Id"ON "Midterm" ("Id");
  ```
  ```python
  >>> !sqlite3 students.db "SELECT * FROM Midterm"
  ```
  ```
  1|Alice|97.3
  2|Bob|87.7
  3|Chris|91.5
  ```

<!-- grade_id: save_answer, points: 0, grade: false, solution: true -->
```python
def save_midterm(df, filename, tablename):
    # YOUR CODE HERE
```

```python
save_midterm(midterm, "students.db", "Midterm")
!sqlite3 students.db ".schema"
```

```python
!sqlite3 students.db "SELECT * FROM Midterm"
```

<!-- grade_id: save_run, points: 5, grade: true, solution: false -->
```python
assert_true(os.path.exists("students.db"))
midterm_test = !sqlite3 students.db "SELECT * FROM Midterm"
print(midterm_test)
midterm_answer = ['1|Alice|97.3', '2|Bob|87.7', '3|Chris|91.5']
assert_equal(len(midterm_test), len(midterm_answer))
assert_equal(set(midterm_test), set(midterm_answer))
```

## Execute an SQL query statement

- Write a function that uses the `read_sql()` function to execute a `SELECT` statement and returns the resulting output as a Pandas DataFrame.
- `save_midterm` takes three arguments: `filename`, `query`, and `index_col`.
- The `filename` parameter is a string and specifies the name of the database file, e.g., `students.db`.
- The `query` parameter is an SQL statement (a string).
- The `index_col` is also a string and specifies the name of the column to set as index, e.g. `"Id"` in the `Midterm` table.
- When we run
  ```python
  >>> data = read_midterm("students.db", query, index_col="Id")
  ```
  the return value `data` should be a Pandas DataFrame. Recall that we wrote data in `midterm` into `students.db`. So, when we use `read_midterm()` (which is basically the reverse of `save_midterm()`), the resulting DataFrame, i.e. `data`, should be identical to the original data.
  ```python
  >>> data = read_midterm("students.db", "SELECT * FROM Midterm;", index_col="Id")
  ```
  ```
       Name  Score
  Id              
  1   Alice   97.3
  2     Bob   87.7
  3   Chris   91.5
   ```

<!-- grade_id: read_answer, points: 0, grade: false, solution: true -->
```python
def read_midterm(filename, query, index_col):
    # YOUR CODE HERE
```

```python
data = read_midterm("students.db", "SELECT * FROM Midterm;", index_col="Id")
print(data)
```

<!-- grade_id: read_test, points: 5, grade: true, solution: false -->
```python
assert_frame_equal(data, midterm)
```

```python
!rm students.db
```

```python

```