If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Cube.

The main purpose of this problem is to help you understand how a typical assignment works.

Let's start with an import statement. In Python, it is considered a good practice to put all import statements at the top. Don't worry about it if you don't understand what the following import statement means. It is just a tool that gives you an immediate feedback on your code. It will also be used in the autograder.

```python
from nose.tools import assert_equal
```

In the following cell is a function definition. Inside the body of this function, we see a comment, `# YOUR CODE HERE`. Without some code between `def cube(number)` and `return result`, this function is incomplete and does nothing. (The red text that starts and ends with """ is called [docstring](https://en.wikipedia.org/wiki/Docstring). It's just a comment that describes what this function is supposed to do, and has no effect.)

Your task is to complete this function by writing code between `def cube(number)` and `return result`.

- Write a function named `cube` that takes a number and returns the cube of that number.

<!-- grade_id: cube_answer, points: 0, grade: false, solution: true -->
```python
def cube(number):
    """
    Computes the cube of "number".
    
    Paramters
    ---------
    number: An integer.
    
    Returns
    -------
    An integer.
    """
    
    # YOUR CODE HERE
    
    return result
```

In the following code cell, we see some assert statements, and every assignment will have such tests. These are what are called [unit tests](https://en.wikipedia.org/wiki/Unit_testing). When you submit your assignment, instructors will determine your assignment grades by using automated grading tools to run these tests on your code, so if you pass all the tests without an error, you get a perfect score (for instructors' grades).

They also give you immediate feedback, because if your code is incorrect, one of the assert statements will produce an error. If the code cell runs without an error, that means your solution is correct.

You don't have to understand what the following code cell means. Just make sure that you don't see any errors when you run it.

<!-- grade_id: cube_test, points: 20, grade: true, solution: false -->
```python
assert_equal(cube(0), 0)
assert_equal(cube(1), 1)
assert_equal(cube(2), 8)
assert_equal(cube(3), 27)
assert_equal(cube(4), 64)
```

This is how pretty much every assignment is going to look like. You are given one or more function definitions, and you write your code where it says `# YOUR CODE HERE` so that all assertments execute without an error.