If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint)

5. You are allowed to submit an assignment multiple times.

# Problem 2. Sum.

In this problem, we will use what we have learned about iterations to write a function that computes a sum of numbers in a sequence.

```python
from nose.tools import assert_equal
```

- Write a function named `sum_from_one_to` that takes an integer and computes the sum of all numbers from 1 to that number. For example, the sum of all numbers from 1 to 4 is 1 + 2 + 3 + 4 = 10, so `sum_from_one_to(4)` should return `10`. In mathematical terms, the `sum_from_one_to(number)` function computes

$$ f(n) = \sum_{i=1}^n n = 1 + 2 + 3 + \cdots + n $$

Hint: Use `for` loops or `while` loops.

<!-- grade_id: sum_from_one_to_answer, points: 0, grade: false, solution: true -->
```python
def sum_from_one_to(number):
    '''
    Computes the sum of all numbers from 1 to "number"
    
    Parameters
    ----------
    number: An integer.
    
    Returns
    -------
    An integer.
    '''
    
    # YOUR CODE HERE
    
    return result
```

You don't have to understand what the following code cell means. Just make sure that you don't see any errors when you run it.

<!-- grade_id: sum_from_one_to_test, points: 30, grade: true, solution: false -->
```python
assert_equal(sum_from_one_to(1), 1)
assert_equal(sum_from_one_to(2), 3)
assert_equal(sum_from_one_to(3), 6)
assert_equal(sum_from_one_to(4), 10)
assert_equal(sum_from_one_to(5), 15)
assert_equal(sum_from_one_to(10), 55)
assert_equal(sum_from_one_to(100), 5050)
```