If you are not using the `Assignments` tab on the course JupyterHub server to read this notebook, read [Activating the assignments tab](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/assignments/README.md).

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 3. Letter Grade.

In this problem, we will use what we have learned about conditionals to write a function that converts a percentage score to a letter grade.

```python
from nose.tools import assert_equal
```

- Write a function that takes a floating point number (i.e. `float`) and returns a letter grade according to the [grading scale](https://github.com/UI-DataScience/accy570-fa16/blob/master/syllabus.md#grading-scale) of ACCY 570. To be precise, `get_letter_grade()` implements the following function:

$$
    f(x) = \begin{cases}
        A, & \text{if } x \geq 90\\
        B, & \text{if } 80 \leq x < 90\\
        C, & \text{if } 70 \leq x < 80\\
        D, & \text{if } 60 \leq x < 70\\
        F, & \text{if } x < 60\\
        \end{cases}
$$

<!-- grade_id: get_letter_grade_answer, points: 0, grade: false, solution: true -->
```python
def get_letter_grade(percent):
    '''
    Assigns a letter grade to a percentage score.
    
    Paramters
    ---------
    percent: A float.
    
    Returns
    -------
    A string, e.g. 'A', 'B', 'C', 'D', or 'F'.
    '''
    
    # YOUR CODE HERE
    
    return grade
```

Run the following code cell to check if your solution is correct. If your solution is correct, you will see no output; otherwise, you will see an error message.

<!-- grade_id: get_letter_grade_test, points: 30, grade: true, solution: false -->
```python
assert_equal(get_letter_grade(101.5), 'A')
assert_equal(get_letter_grade(100.0), 'A')
assert_equal(get_letter_grade(95.0), 'A')
assert_equal(get_letter_grade(90.0), 'A')

assert_equal(get_letter_grade(89.9), 'B')
assert_equal(get_letter_grade(84.5), 'B')
assert_equal(get_letter_grade(80.0), 'B')

assert_equal(get_letter_grade(79.9), 'C')
assert_equal(get_letter_grade(72.1), 'C')
assert_equal(get_letter_grade(70.0), 'C')

assert_equal(get_letter_grade(69.9), 'D')
assert_equal(get_letter_grade(67.8), 'D')
assert_equal(get_letter_grade(60.0), 'D')

assert_equal(get_letter_grade(59.9), 'F')
assert_equal(get_letter_grade(42.8), 'F')
assert_equal(get_letter_grade(34.2), 'F')
assert_equal(get_letter_grade(11.0), 'F')
assert_equal(get_letter_grade(0.0), 'F')
```