A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 2. Functional Programming.

In this problem, we will apply what we have learned about the `map`, `filter`, and `reduce` methods, and lambda functions.

```python
from functools import reduce

from nose.tools import assert_equal
```

## Map

- Write a function named `scale_list` that scales each number in a list by a given factor.
- The function takes two arguments: `items` is a list, and `factor` is an int.
- The function returns a list where each element in `items` is scaled by `factor`.
- For example,

```python
>>> print(scale_list([1, 2, 3, 4, 5], 10))
```
```
[10, 20, 30, 40, 50]
```

- We are using **functional programming** in this problem. In imperative programming, this function would look something like

```python
def scale_list(items, factor):
    result = []
    for i in items:
        result.append(i * factor)
    return result
```

In functional programming, we don't have side effects. That means, no `for` loops in functional programming. Thus, your function in this problem should not have a `for` loop. Use a combination of `map`, `list`, and `lambda`.

- When I say no `for` loops, I don't mean you can do

```python
def scale_list(items, factor):
    return items * factor
```

The goal of this problem is to give you practice in using the `map` function. Use a combination of `map`, `list`, and `lambda`. The autograder (for instructor assessment) won't know the difference even if you use a `for` loop or simply use the `*` operator on the list. Therefore, **when you are peer grading, you should consider the solution incorrect if the student did not use the `map` function.**

<!-- grade_id: scale_list_answer, points: 0, grade: false, solution: true -->
```python
def scale_list(items, factor):
    # YOUR CODE HERE
```

```python
print(scale_list([1, 2, 3, 4, 5], 10))
```

<!-- grade_id: scale_list_test, points: 5, grade: true, solution: false -->
```python
assert_equal(scale_list([1, 2, 3, 4, 5], 10), [10, 20, 30, 40, 50])
assert_equal(scale_list([1, 2, 3, 4, 5], 0), [0, 0, 0, 0, 0])
assert_equal(scale_list([1, 2, 3, 4, 5], 1), [1, 2, 3, 4, 5])

assert_equal(scale_list([0, 1, 2], 5), [0, 5, 10])
assert_equal(scale_list([0, 1, 2], 0), [0, 0, 0])
assert_equal(scale_list([0, 1, 2], 1), [0, 1, 2])
```

## Reduce

- Write a function named `accumulate` that computes the sum of all number in a list.
- The function takes one argument `items`. `item` is a list.
- The function returns an integer.
- For example,

```python
>>> print(accumulate([1, 2, 3, 4, 5]))
```
```
15
```

- In imperative programming, this function would look something like

```python
def accumulate(items):
    result = 0
    for i in items:
        result += i
    return result
```

However, you should not use a `for` loop in this problem. Use a combination of `reduce` and `lambda`.

- When I say no `for` loops, I don't mean you can do

```python
def accumulate(items):
    return sum(items)
```

The goal of this problem is to give you practice in using the `reduce` function. Use a combination of `reduce` and `lambda`. The autograder (for instructor assessment) won't know the difference even if you use a `for` loop or simply use the `sum` function. Therefore, **when you are peer grading, you should consider the solution incorrect if the student did not use the `reduce` function.**

<!-- grade_id: accumulate_answer, points: 0, grade: false, solution: true -->
```python
def accumulate(items):
    # YOUR CODE HERE
```

```python
print(accumulate([1, 2, 3, 4, 5]))
```

<!-- grade_id: accumulate_test, points: 5, grade: true, solution: false -->
```python
assert_equal(accumulate([1]), 1)
assert_equal(accumulate([1, 2]), 3)
assert_equal(accumulate([1, 2, 3]), 6)
assert_equal(accumulate([1, 2, 3, 4]), 10)
assert_equal(accumulate([1, 2, 3, 4, 5]), 15)
```

## Filter

- Write a function named `filter_odd` that selects only the odd numbers in a list.
- The function takes one argument `items`. `item` is a list.
- The function returns a list.
- For example,

```python
>>> print(filter_odd([1, 2, 3, 4, 5]))
```
```
[1, 3, 5]
```

- In imperative programming, this function would something like

```python
def filter_odd(items):
    result = []
    for i in items:
        if i % 2:
            result.append(i)
    return result
```

However, you should not use a `for` loop in this problem. Use a combination of `reduce` and `lambda`.

- Many people would argue that the following function is a better solution than using `filter`:

```python
def filter_odd(items):
    return [i for i in items if i % 2]
```

However, the goal of this problem is to give you practice in using the `filter` function. Use a combination of `list`, `filter`, and `lambda`. The autograder (for instructor assessment) won't know the difference even if you use a `for` loop or list comprehensions. Therefore, **when you are peer grading, you should consider the solution incorrect if the student did not use the `filter` function.**

<!-- grade_id: filter_odd_answer, points: 0, grade: false, solution: true -->
```python
def filter_odd(items):
    # YOUR CODE HERE
```

```python
print(filter_odd([1, 2, 3, 4, 5]))
```

<!-- grade_id: filter_odd_test, points: 5, grade: true, solution: false -->
```python
assert_equal(filter_odd([1, 2, 3, 4, 5]), [1, 3, 5])
assert_equal(filter_odd([1, 2, 3]), [1, 3])
assert_equal(filter_odd([1, 2, 3, 5]), [1, 3, 5])
assert_equal(filter_odd([2, 4]), [])
```

```python

```