A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Unix Networking

In this problem, we use Unix commands to download and extract a compressed file.

```python
import os
from nose.tools import assert_equal, assert_false
```

We have been using the [airline on-time performance data](http://stat-computing.org/dataexpo/2009/the-data.html) in lessons and assignments. Since we are using a centralized Jupyterhub server in this course, this file was provided for you. However, suppose you are working in a different environment, and you have to download the data. How would you use Unix commands to download a file, and if the file is compressed, extract it so you can perform some text processing?

## Use `wget` to download the [airline on-time performance data](http://stat-computing.org/dataexpo/2009/the-data.html) to the current directory.

- Use the 2001 data. The url is http://stat-computing.org/dataexpo/2009/2001.csv.bz2.
- Note the file name ends with `.bz2`, a common file extension for [bzip2](https://en.wikipedia.org/wiki/Bzip2) compressed files.
- Don't forget that you need the exclamation point (`!`) to run Unix commands in Jupyter notebooks.
- See the [Unix Networking notebook](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week10/notebooks/unixnetwork.ipynb) and Chapter 16 of [The Linux Command Line](https://sourceforge.net/projects/linuxcommand/?source=dlp).

<!-- grade_id: wget_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

You might want to use `ls` to list the contents of the current working directory and see if there is `2001.csv.bz2`.

```python
!ls
```

In the following code cell, we will calculate the [checksum](https://en.wikipedia.org/wiki/Checksum) of the file to check if you have the correct file. Don't worry if you don't understand the code (or what checksum is); just make sure that it doesn't produce any errors.

<!-- grade_id: wget_test, points: 5, grade: true, solution: false -->
```python
def get_checksum(filename):
    import hashlib
    with open(filename, 'rb') as f:
        md5 = hashlib.md5(f.read())
    return md5.hexdigest()

assert_equal(get_checksum('2001.csv.bz2'), 'e855ba7cc04f560199953833305a0f90')
```

Since the downloaded file is compressed, we have to extract it.

## Extract the compressed file.

- Use `bzip2` with the `-d` option to extract the flight data file you downloaded in the previous section.
- See Chapter 18 of [The Linux Command Line](https://sourceforge.net/projects/linuxcommand/?source=dlp).

<!-- grade_id: bzip2_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

You might want to use `ls` to list the contents of the current working directory and see if the file was extracted successfully to `2001.csv`.

```python
!ls
```

The following code cells are for testing your answer and performing clean-up.

<!-- grade_id: bzip2_test, points: 5, grade: true, solution: false -->
```python
assert_equal(get_checksum('2001.csv'), '3ca459bb83ad2074fcb4c8b88d28fcaf')
assert_false(os.path.exists('2001.csv.bz2'))
```

```python
!rm $(pwd)/2001.csv*
```

```python

```