A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 2. Unix Text Processing.

In this problem, we use Unix commands to perform some basic text processing.

```python
from nose.tools import assert_equal
```

- Use `sed` to change the commas in the `2001.csv` file to vertical bars, and save the result to a file named `2001.csv.sed`.

- Don't forget that you need the exclamation point (`!`) to run Unix commands in Jupyter notebooks.

- Hint: We use the substitute pattern `s/../../` to convert some pattern to another pattern. For example, suppose we have a CSV file named `input.csv` that looks as follows:

```shell
$ head input.csv
```

```
Month,DayofMonth,ArrTime,Dest
1,17,1931,CLT
1,18,1938,CLT
1,19,1957,CLT
1,20,1944,CLT
1,21,1954,CLT
1,22,1931,CLT
1,23,1924,CLT
1,24,1922,CLT
1,25,1925,CLT
```

The `head` command prints out the first 10 lines by default. (The number of lines can be changed with the `-n` parameter.) To change commas in `input.csv` to vertical bars and save the results as `output.bsv`, we would type

```shell
$ sed 's/,/|/' <input.csv >output.bsv
```

Note the less-than symbol (`<`) redirects a file (`input.csv` in this case) as _input_ for a command (`sed`). The greater-than symbol (`>`) redirects the _output_ of a command to a file (`output.bsv`). Let's look at what the resulting file looks like.

```shell
$ head output.bsv
```

```
Month|DayofMonth,ArrTime,Dest
1|17,1931,CLT
1|18,1938,CLT
1|19,1957,CLT
1|20,1944,CLT
1|21,1954,CLT
1|22,1931,CLT
1|23,1924,CLT
1|24,1922,CLT
1|25,1925,CLT
```

Note that only the first occurrences of the commas are changed on every line. This behavior is typical of sed; it looks at one line at a time, and changes only the first occurrence of the search word. Here, we want to make changes on every occurrence instead of only the first. The `g` (global) flag at the end tells sed to make changes for every occurrence. That is, we have to use the command `s/../../g`.

- Note, the `2001.csv` file at `/home/data_scientist/data` is a pretty big file (about 6 million lines, 573 megabytes), so it may take a while (30-40 seconds) to complete.

- Type your answer in the following code cell.

<!-- grade_id: sed_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

You might want to print out the first few lines, and visually inspect the output to check if all commas were converted to vertical bars.

```python
!head 2001.csv.sed
```

The following code cells are for testing your answer and performing clean-up.

<!-- grade_id: sed_test, points: 10, grade: true, solution: false -->
```python
sed_head = !cat 2001.csv.sed | head
length = !wc -l /home/data_scientist/data/2001.csv

for i in range(10):
    assert_equal(len(sed_head[i].split('|')), 29)
    
assert_equal(sed_head[0].split('|')[:5], ['Year', 'Month', 'DayofMonth', 'DayOfWeek', 'DepTime'])
assert_equal(sed_head[1].split('|')[:5], ['2001', '1', '17', '3', '1806'])
assert_equal(length.s.split()[0], '5967781')
```

```python
!rm $(pwd)/2001.csv*
```

```python

```