A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 2.  SQL Schema

In this problem, we use the SQLite database to build a fictitious database.

```python
import os
from nose.tools import assert_equal, assert_true
```

Suppose we have the following database table:

**<DIV ALIGN="CENTER"> Students </DIV>**

| Id | Name |
| -- | ---- |
| 1  | Alice|
| 2  | Bob  |
| 3  | Chris |

We want to create this table in SQLite. As [The Introduction to SQL notebook](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week11/notebook/intro2sqlddl.ipynb) explains, we need to create an SQL schema. We will follow the lesson notebook to write our schema as a plain text file, and then use the `sqlite3` tool at the Unix command prompt.

## Create a schema file

- In the following code cell, create a schema file for the `Students` table with two columns, `Id` and `Name`.

- The name of your table must be `Students`.

- We use the [%%writefile magic command](http://ipython.readthedocs.io/en/stable/interactive/magics.html#cellmagic-writefile) to write `students.sql`. Anything you write below the first line (`%%writefile students.sql`) will be saved to a file named `students.sql`.

- To indicate a comment in Python, we use `#`. However, SQL uses a different symbol to indicate a comment. If you want to comment your code, use two hyphens (`--`). Any line that starts with `--` will be ignored by SQLite.

<!-- grade_id: writefile_answer, points: 0, grade: false, solution: true -->
```python
%%writefile students.sql

-- YOUR CODE HERE

```

## Create the schema in a new database

- In the following code cell, write a Unix command that creates the schema in a new database named `students.db`.

<!-- grade_id: sqlite3_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

We can print out the results using the `.schema` command in `sqlite3`. The output of
```shell
$ sqlite3 students.db ".schema"
```
should match the schema in the `students.sql` file.

```python
!sqlite3 students.db ".schema"
```

<!-- grade_id: sqlite3_test, points: 20, grade: true, solution: false -->
```python
assert_true(os.path.exists("students.db"))

!sqlite3 students.db "DELETE FROM students"
!sqlite3 students.db "INSERT INTO students VALUES (1, 'Alice')"
!sqlite3 students.db "INSERT INTO students VALUES (2, 'Bob')"
!sqlite3 students.db "INSERT INTO students VALUES (3, 'Chris')"
test_result = !sqlite3 students.db "SELECT * FROM Students"
assert_equal(test_result, ['1|Alice', '2|Bob', '3|Chris'])
```

```python
!rm students.*
```

```python

```