A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. Data Persistence

In this problem, we use pickling to save data to a file and to later reconstitute the data.

```python
import os
import pickle

from nose.tools import assert_true, assert_equal, assert_is_instance
```

[The Introduction to Data Persistence](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week11/notebook/intro2db.ipynb) notebook shows an example of pickling and unpickling a numpy array. But it's not just arrays (or ints, floats, strings, lists, dictionaries) that can be pickled. Just about anything in Python can be pickled &mdash; even functions!

In the following code cell is a simple function that calculates the square of a number.

```python
def square(number):
    return number**2
```

```python
for i in range(10):
    print("The square of {0} is {1}.".format(i, square(i)))
```

## Pickle data and save to a file.

- Write a function named `save` that accepts two arguments: `data` is data (or a function in our case) that will be pickled, and `filename` is a string.
- For example, to save the `sqaure()` function to a file named `square.pickle`, we would write as follows:
```python
>>> save(square, "square.pickle")
```
(Note the lack of parentheses after `square`. Parentheses are used when we want to _call_ a function. Here, we are not calling the function; we are passing the function itself as an argument.)
- In this problem, we use `save()` to save a function, but `save()` can be used to save any other data structures, such as numpy arrays, strings, lists, dictionaries, etc. because code will be identical for all cases.

<!-- grade_id: save_answer, points: 0, grade: false, solution: true -->
```python
def save(data, filename):
    """
    Uses pickle.dump() to pickle 'data'
    and saves it to a file specified by 'filename'.
    
    Parameters
    ----------
    data: Any pickable object.
    filename: A string
    
    Returns
    -------
    None
    """
    
    # YOUR CODE HERE
```

In the following code cell, we pickle the `square` function and save it as `square.pickle`. Again, note the lack of parentheses after `square`.

```python
save(square, "square.pickle")
```

<!-- grade_id: save_test, points: 10, grade: true, solution: false -->
```python
assert_true(os.path.exists("square.pickle"))

def get_checksum(filename):
    import hashlib
    with open(filename, 'rb') as f:
        md5 = hashlib.md5(f.read())
    return md5.hexdigest()

assert_equal(
    get_checksum("square.pickle"),
    "8ba2f8b9d821cab352629fc4bf54d6e7"
)

# test one more case
fname = "save.pickle"
save(save, fname)
assert_equal(get_checksum(fname), "5638258fc985a2c4e97f1b33d31a0aa3")
os.remove(fname)
```

## Unpickle data.

- Write a function named `load` that accepts a `filename`, unpickles the file, and returns the unpickled data.
- For example, to load the data saved in `square.pickle`, we would write as follows:
```python
>>> unpickled_square = load("square.pickle")
```
- We are using `load()` to unpickle a function, but `load()` can be used to load any other data structure, such as numpy arrays, strings, lists, dictionaries, etc. because code will be identical for all cases.

<!-- grade_id: load_answer, points: 0, grade: false, solution: true -->
```python
def load(filename):
    """
    Unpickles data saved in 'filename'.
    
    Parameters
    ----------
    filename: A string
    
    Returns
    -------
    A function (but in principle, can be any Python object.)
    """
    
    # YOUR CODE HERE
```

```python
unpickled_square = load("square.pickle")
```

After unpickling the function, we can use it in the same way that we use `square()`.

```python
for i in range(10):
    print("The square of {0} is {1}.".format(i, unpickled_square(i)))
```

<!-- grade_id: load_test, points: 10, grade: true, solution: false -->
```python
assert_true(callable(unpickled_square))
for i in range(100):
    assert_equal(unpickled_square(i), square(i))
    
# test one more case
fname = "test.pickle"
def cube(x):
    return x**3
save(cube, fname)
unpickled_cube = load(fname)
for i in range(100):
    assert_equal(unpickled_cube(i), cube(i))
os.remove(fname)
```

```python
os.remove("square.pickle")
```

```python

```