A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 1. List

The goal of this problem is to get you more comfortable with the Python data structures, such as string, tuple, list, and dictionary. We focus on the _list_ in this problem, and we will focus on the _dictionary_ in Problem 2.

Let's start with an import statement. Again, you don't have to know what the `nose` library does. It is just a libarry with useful testing tools.

```python
from nose.tools import assert_equal, assert_true, assert_is_instance
```

## Write a function that returns a list of all odd numbers between 0 and 10 in increasing order, i.e. 1, 3, 5, 7, and 9.

There are several ways to create a list. As long as your function returns the correct list and the code cell with tests (`assert...`) does not produce any errors, all methods are correct and there is no wrong way to do it (but it can be argued that some methods are more concise and/or more efficient than others). To create the list, you can use any of the following approaches:

- You can always type in the values by hand and enclose them with square brackets.
- You can also use a regular `for` loop as follows. Start with an empty list, use a loop to go from 0 to 10, check if the current number is an odd integer, and if so, append to the list.
- You can use list comprehensions. List comprehensions were covered in the [Flow Control notebook](https://github.com/UI-DataScience/accy570-fa16/blob/master/Week1/notebooks/flowcontrolpy.ipynb) and also in [Dive into Python 3](http://www.diveintopython3.net/comprehensions.html).
- You can use a combination of `list()` and `range()`.
- Or, any other methods that produce the desired result.

Sample output:

```python
>>> odd_integers = get_some_odd_integers()
>>> print(odd_integers)
```
```
[1, 3, 5, 7, 9]
```

<!-- grade_id: get_some_odd_integers_answer, points: 0, grade: false, solution: true -->
```python
def get_some_odd_integers():
    """
    Returns a list of all odd numbers between 0 and 10 in increasing order.
    That is, this function always returns the list [1, 3, 5, 7, 9].
    """
    
    # YOUR CODE HERE
    
    return result
```

The easiest way to check if the output is correct is to simply print out the result. The following code cell stores the output of the `get_some_odd_integers()` function in a variable named `odd_integers`, and we visually check if our function is correct by printing it on the screen.

```python
odd_integers = get_some_odd_integers()
print(odd_integers)
```

Test your function by running the following code cell. If the code cell runs without an error, that means your solution is correct.

As always, you don't have to understand what `assert_is_instance`, `assert_true`, etc. mean. Just make sure that you don't see any errors when you run it, and if it gives you an error, use the error message as a clue to fix your code.

<!-- grade_id: get_some_odd_integers_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(odd_integers, list)
assert_true(all([isinstance(i, int) for i in odd_integers]))
assert_true(0 not in odd_integers)
assert_true(1 in odd_integers)
assert_true(2 not in odd_integers)
assert_true(3 in odd_integers)
assert_true(4 not in odd_integers)
assert_true(5 in odd_integers)
assert_true(6 not in odd_integers)
assert_true(7 in odd_integers)
assert_true(8 not in odd_integers)
assert_true(9 in odd_integers)
assert_true(10 not in odd_integers)
```

## Write a function named `tail` that returns the last value of a list.

For example, when we use the `tail` function on `[1, 3, 5, 7, 9]`, it should return `9`, and when we have `[2, 3, 5, 7, 11]`, it should return `11`.

Sample output:
```python
>>> data = [2, 3, 5, 7, 11]
>>> print(tail(data))
```
```
11
```

<!-- grade_id: tail_answer, points: 0, grade: false, solution: true -->
```python
def tail(a_list):
    """
    Takes a list and returns the last value.
    
    Parameters
    ----------
    a_list: An arbitrary list of integers.
    
    Returns
    -------
    An integer. The last value of 'a_list'.
    """
    
    # YOUR CODE HERE
    
    return result
```

```python
print(tail(odd_integers))
```

<!-- grade_id: tail_test, points: 10, grade: true, solution: false -->
```python
assert_equal(tail(odd_integers), 9)

# test more cases
assert_equal(tail([2, 3, 5, 7, 11]), 11)
assert_equal(tail([2, 3, 5, 7, 11, 13]), 13)
```

## Write a function named `head` that returns a list with all elements of the given list _except_ the last element.

For example, when we use the `head` function on `[1, 3, 5, 7, 9]`, it should return `[1, 3, 5, 7]` (we don't want the `9`). And when we have `[2, 3, 5, 7, 11]`, it should return `[2, 3, 5, 7]` (with no `11`).

Sample output:
```python
>>> data = [2, 3, 5, 7, 11]
>>> print(head(data))
```
```
[2, 3, 5, 7]
```

<!-- grade_id: head_answer, points: 0, grade: false, solution: true -->
```python
def head(a_list):
    """
    Returns a list that has all elements of 'a_list' but not the last element.
    
    Parameters
    ----------
    a_list: An arbitrary list of integers.
    
    Returns
    -------
    A list of length, len(a_list) - 1.
    """
    
    # YOUR CODE HERE
    
    return result
```

```python
print(head(odd_integers))
```

<!-- grade_id: head_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(head(odd_integers), list)
assert_true(1 in head(odd_integers))
assert_true(3 in head(odd_integers))
assert_true(5 in head(odd_integers))
assert_true(7 in head(odd_integers))
assert_true(9 not in head(odd_integers))

# test one more case
assert_equal(head([2, 3, 5, 7, 11]), [2, 3, 5, 7])
assert_equal(head([2, 3, 5, 7, 11, 13]), [2, 3, 5, 7, 11])
```