A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 4. Numpy.

In this problem, we will repeat the same tasks in Problem 1, but we will use Numpy arrays instead of lists.

```python
import numpy as np
from nose.tools import assert_equal, assert_is_instance
from numpy.testing import assert_array_equal
```

## Write a function that returns a Numpy array of all odd numbers between 0 and 10 in increasing order, i.e. 1, 3, 5, 7, and 9.

There are several ways to do this. Again, as long as your function returns the correct array and the code cell with tests (`assert...`) does not produce any errors, all methods are correct and there is no wrong way to do it. To create the array, you can use any of the following approaches:

- You can create a Python list and convert it to a numpy array using `np.array()`.
- You can use `np.arange()` and specify the step size.
- You can also use `np.linspace()`.
- Or, any other methods that produce the desired result.

Sample output:

```python
>>> odd_integers = get_some_odd_integers()
>>> print(odd_integers)
```
```
[1 3 5 7 9]
```
```python
>>> print(type(odd_integers))
```
```
<class 'numpy.ndarray'>
```

<!-- grade_id: get_some_odd_integers_answer, points: 0, grade: false, solution: true -->
```python
def get_some_odd_integers():
    """
    Returns a Numpy array of all odd numbers between 0 and 10 in increasing order.
    That is, this function always returns the array [1, 3, 5, 7, 9].
    """
    
    # YOUR CODE HERE
    
    return result
```

```python
odd_integers = get_some_odd_integers()
print(odd_integers)
print(type(odd_integers))
```

<!-- grade_id: get_some_odd_integers_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(odd_integers, np.ndarray)
assert_array_equal(odd_integers, [1, 3, 5, 7, 9])
```

## Write a function named `tail` that returns the last value of an array.

For example, when we use the `tail` function on the array `[1 3 5 7 9]`, it should return `9`, and when we have the array `[2 3 5 7 11]`, it should return `11`.

Sample output:
```python
>>> print(tail(odd_integers))
```
```
9
```

<!-- grade_id: tail_answer, points: 0, grade: false, solution: true -->
```python
def tail(array):
    """
    Takes an array and returns the last value.
    
    Parameters
    ----------
    array: An arbitrary array of integers.
    
    Returns
    -------
    An integer. The last value of 'array'.
    """
    
    # YOUR CODE HERE
    
    return result
```

```python
print(tail(odd_integers))
```

<!-- grade_id: tail_test, points: 10, grade: true, solution: false -->
```python
assert_equal(tail(odd_integers), 9)

# test more cases
assert_equal(tail(np.array([2, 3, 5, 7, 11])), 11)
assert_equal(tail(np.array([2, 3, 5, 7, 11, 13])), 13)
```

## Write a function named `head` that returns an array with all elements of the given array _except_ the last element.

For example, when we use the `head` function on `[1 3 5 7 9]`, it should return `[1 3 5 7]` (we don't want the `9`). And when we have `[2 3 5 7 11]`, it should return `[2 3 5 7]` (with no `11`).

Sample output:
```python
>>> print(head(odd_integers))
```
```
[1 3 5 7]
```

<!-- grade_id: head_answer, points: 0, grade: false, solution: true -->
```python
def head(array):
    """
    Returns an array that has all elements of 'array' but not the last element.
    
    Parameters
    ----------
    array: An arbitrary array of integers.
    
    Returns
    -------
    An array of length, len(array) - 1.
    """
    
    # YOUR CODE HERE
    
    return result
```

```python
print(head(odd_integers))
```

<!-- grade_id: head_test, points: 10, grade: true, solution: false -->
```python
assert_is_instance(head(odd_integers), np.ndarray)
assert_array_equal(head(odd_integers), [1, 3, 5, 7])

# test one more case
assert_array_equal(head(np.array([2, 3, 5, 7, 11])), [2, 3, 5, 7])
assert_array_equal(head(np.array([2, 3, 5, 7, 11, 13])), [2, 3, 5, 7, 11])
```