A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer in anywhere else other than where it says `YOUR CODE HERE`. Anything you write anywhere else will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_)

5. You are allowed to submit an assignment multiple times, but only the most recent submission will be graded.

# Problem 2. Dictionary.

The goal of this problem is to get you more comfortable with the Python data structures, such as string, tuple, list, and dictionary. We looked at the list in Problem 1. We focus on the _dictionary_ in this problem.

```python
from nose.tools import assert_equal, assert_is_instance, assert_true
```

## Write a function that creates a dictionary whose values are as follows.

```
"banana": 5
"apple": 1.5
"orange": 2
"pear": 4
```

`"banana"`, `"apple"`, etc. are the keys of the dictionary, and the numbers, `5`, `1.5`, etc., are the values.

<!-- grade_id: get_a_dict_of_prices_answer, points: 0, grade: false, solution: true -->
```python
def get_a_dict_of_prices():
    """
    Returns a dictionary with prices of 4 fruits.
    Banana = $4, apple = $1.5, orange = $2, pear = $4.
    """
    
    # YOUR CODE HERE
    
    return prices
```

```python
prices = get_a_dict_of_prices()
print(prices)
```

<!-- grade_id: get_a_dict_of_prices_test1, points: 10, grade: true, solution: false -->
```python
assert_is_instance(prices, dict)

assert_equal(len(prices.keys()), 4)
assert_true("banana" in prices)
assert_true("apple" in prices)
assert_true("orange" in prices)
assert_true("pear" in prices)

assert_equal(len(prices.values()), 4)
assert_equal(set(prices.values()), set([5, 1.5, 2, 4]))
```

## Write a function that takes a dictionary and returns the price of banana.

Assume that `banana` is always one of the keys of the dictionary.

Sample output:

```python
>>> print(get_price_of_banana(prices))
```
```
5
```

<!-- grade_id: get_price_of_banana_answer, points: 0, grade: false, solution: true -->
```python
def get_price_of_banana(a_dictionary):
    """
    Returns the price of banana. 'a_dictionary' has the key 'banana'.
    
    Parameters
    ----------
    a_dictionary: A dict. One of its keys is 'banana'.
    
    Returns
    -------
    An int. Price of banana.
    """
    
    # YOUR CODE HERE
    
    return price
```

```python
print(get_price_of_banana(prices))
```

<!-- grade_id: get_price_of_banana_test, points: 10, grade: true, solution: false -->
```python
assert_equal(get_price_of_banana(prices), 5)

# test some more cases
assert_equal(get_price_of_banana({"banana": 6, "placeholder1": 1}), 6)
assert_equal(get_price_of_banana({"banana": 7, "placeholder2": 2, "placeholder3": 3}), 7)
```

Now the `get_price_of_banana` function that we have just written will work only on bananas. We also have apples, oranges, and pears. But using the same approach to get their prices, we will have to write three more functions (that look very similar). If we think about it, the keys (i.e. `banana`, `apple`, `orange`, and `pear`) are variables, and we can combine the four cases and generalize the previous function to handle any arbitrary keys, by using another parameter in the function.

## Write a function that takes _two_ parameters, `prices` (a dictionary) and `item` (a string), and returns the price of the item.

Assume that the item you are searching for will always be in the dictionary.

Sample output:

```python
>>> print(get_price_of_item(prices, "banana"))
```
```
5
```
```python
>>> print(get_price_of_item(prices, "apple"))
```
```
1.5
```
```python
>>> print(get_price_of_item(prices, "orange"))
```
```
2
```
```python
>>> print(get_price_of_item(prices, "pear"))
```
```
4
```

<!-- grade_id: get_price_of_item_answer, points: 0, grade: false, solution: true -->
```python
def get_price_of_item(a_dictionary, item):
    """
    Returns the value (price) of 'item' (fruit).
    
    Parameters
    ----------
    a_dictionary: A dict.
    
    Returns
    -------
    An int. 
    """
    
    # YOUR CODE HERE
    
    return price
```

```python
print(get_price_of_item(prices, "banana"))
print(get_price_of_item(prices, "apple"))
print(get_price_of_item(prices, "orange"))
print(get_price_of_item(prices, "pear"))
```

<!-- grade_id: get_price_of_item_test, points: 10, grade: true, solution: false -->
```python
assert_equal(get_price_of_item(prices, "banana"), 5)
assert_equal(get_price_of_item(prices, "apple"), 1.5)
assert_equal(get_price_of_item(prices, "orange"), 2)
assert_equal(get_price_of_item(prices, "pear"), 4)

### test some more cases
assert_equal(get_price_of_item({"banana": 6, "apple": 1}, "banana"), 6)
assert_equal(get_price_of_item({"banana": 6, "apple": 1}, "apple"), 1)
assert_equal(get_price_of_item({"banana": 7, "apple": 2, "orange": 3}, "banana"), 7)
assert_equal(get_price_of_item({"banana": 7, "apple": 2, "orange": 3}, "apple"), 2)
assert_equal(get_price_of_item({"banana": 7, "apple": 2, "orange": 3}, "orange"), 3)
```